SARS-CoV-2 nsp16 is regulated by host E3 ubiquitin ligases, UBR5 and MARCHF7

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

In this study, Tian et al. explore the role of ubiquitination of non-structural protein 16 (nsp16) in the SARS-CoV-2 life cycle. nsp16, in conjunction with nsp10, performs the final step of viral mRNA capping through its 2'-O-methylase activity. This modification allows the virus to evade host immune responses and protects its mRNA from degradation. The authors demonstrate that nsp16 undergoes ubiquitination and subsequent degradation by the host E3 ubiquitin ligases UBR5 and MARCHF7 via the ubiquitin-proteasome system (UPS). Specifically, UBR5 and MARCHF7 mediate nsp16 degradation through K48- and K27-linked ubiquitination, respectively. Notably, degradation of nsp16 by either UBR5 or MARCHF7 operates independently, with both mechanisms effectively inhibiting SARS-CoV-2 replication in vitro and in vivo. Furthermore, UBR5 and MARCHF7 exhibit broad-spectrum antiviral activity by targeting nsp16 variants from various SARS-CoV-2 strains. This research advances our understanding of how nsp16 ubiquitination impacts viral replication and highlights potential targets for developing broadly effective antiviral therapies.

Strengths:

The proposed study is of significant interest to the virology community because it aims to elucidate the biological role of ubiquitination in coronavirus proteins and its impact on the viral life cycle. Understanding these mechanisms will address broadly applicable questions about coronavirus biology and enhance our overall knowledge of ubiquitination's diverse functions in cell biology. Employing in vivo studies is a strength.

Weaknesses:

Minor comments:

Figure 5A- The authors should ensure that the figure is properly labeled to clearly distinguish between the IP (Immunoprecipitation) panel and the input panel.

Thank you for your suggestion. We have exchanged Figure 5 in this version.

Reviewer #3 (Public review):

Summary:

The manuscript "SARS-CoV-2 nsp16 is regulated by host E3 ubiquitin ligases, UBR5 and MARCHF7" is an interesting work by Tian et al. describing the degradation/ stability of NSP16 of SARS CoV2 via K48 and K27-linked Ubiquitination and proteasomal degradation. The authors have demonstrated that UBR5 and MARCHF7, an E3 ubiquitin ligase bring about the ubiquitination of NSP16. The concept, and experimental approach to prove the hypothesis looks ok. The in vivo data looks ok with the controls. Overall, the manuscript is good.

Strengths:

The study identified important E3 ligases (MARCHF7 and UBR5) that can ubiquitinate NSP16, an important viral factor.

Comments on revisions:

I had gone through the revised form of the manuscript thoroughly. The authors have addressed all of my concerns. To me, the experimental approach looks convincing that the host E3 ubiquitin ligases (UBR5 and MARCHF7) ubiquitinate NSP16 and mark it for proteasomal degradation via K48- and K27- linkage. The authors have represented the final figure (Fig.8) in a convincing manner, opening a new window to explore the mechanism of capping the vRNA bu NSP16.

Thank you for your recognition.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Severe acute respiratory syndrome coronavirus 2 (SARS-CoV-2), the virus responsible for coronavirus disease 2019 (COVID-19), remains a global public health threat with major economic implications. The non-structural protein 16 (nsp16) of SARS-CoV-2, in complex with nsp10, catalyses the final step of viral mRNA capping via its 2’-O-methylase activity, enabling the virus to evade host immunity and protect viral mRNA from degradation. However, host factors regulating nsp16 have not been thoroughly explored. Although various E3 ubiquitin ligases are known to interact with SARS-CoV-2 proteins, their specific roles in targeting nsp16 for degradation remain unclear. In this study, we demonstrate that nsp16 is ubiquitinated and degraded by host E3 ubiquitin ligases UBR5 and MARCHF7, acting through the ubiquitin-proteasome system (UPS). UBR5 and MARCHF7 induce nsp16 degradation via K48- and K27-linked ubiquitination, respectively, in independent processes that inhibit SARS-CoV-2 replication both in vitro and in vivo. Furthermore, UBR5 and MARCHF7 exhibited broad-spectrum antiviral activity by degrading nsp16 variants from different SARS-CoV-2 strains. Our findings uncover novel mechanisms by which the UPS antagonises SARS-CoV-2 replication and provide promising targets for therapeutic intervention against COVID-19.

Revised version: Severe acute respiratory syndrome coronavirus 2 (SARS-CoV-2), the causative agent of coronavirus disease 2019 (COVID-19), remains a global public health threat with considerable economic consequences. The non-structural protein 16 (nsp16), in complex with nsp10, facilitates the final viral mRNA capping step through its 2′-O-methylase activity, helping the virus to evade host immunity and prevent mRNA degradation. However, nsp16 regulation by host factors remains poorly understood. While various E3 ubiquitin ligases interact with SARS-CoV-2 proteins, their roles in targeting nsp16 for degradation are unclear. In this study, we demonstrate that nsp16 undergoes ubiquitination and proteasomal degradation mediated by the host E3 ligases UBR5 and MARCHF7. UBR5 induces K48-linked ubiquitination, whereas MARCHF7 promotes K27-linked ubiquitination, independently suppressing SARS-CoV-2 replication in vitro and in vivo. Notably, UBR5 and MARCHF7 also degrade nsp16 variants from different viral strains, exhibiting broad-spectrum antiviral activity. Our findings reveal novel antiviral mechanisms of the ubiquitin-proteasome system (UPS) and highlight their potential therapeutic targets against COVID-19.

Introduction

Earlier version: Severe acute respiratory syndrome coronavirus 2 (SARS-CoV-2) has emerged as a global public health threat, causing over 700 million coronavirus disease 2019 (COVID-19) cases and 7 million deaths (Carabelli et al., 2023; Markov et al., 2023). SARS-CoV-2 encodes four structural proteins, 16 non-structural proteins (Nsps) that assist in viral replication and transcription, and a series of accessory proteins (ORFs) associated with immune evasion (V’kovski et al., 2021; Wang et al., 2020). Some viruses such as Ebola virus (Valle et al., 2021), dengue virus (Jung et al., 2018), and reovirus (Furuichi et al., 1975) encode 2’-O methyltransferases (2’-O-MTases) homologous to those found in eukaryotes. This allows them to mimic the host 5’ cap structure on their RNA, hindering recognition by the innate immune system (Daffis et al., 2010; Züst et al., 2011). Similarly, SARS-CoV-2 nsp16, in complex with nsp10, functions as a 2’-O-MTase, methylating the first nucleotide of a capped viral RNA strand at the 2′-O position, converting the 5’ cap of RNA from ‘cap-0’ to ‘cap-1’ (Benoni et al., 2021; Park et al., 2022). This modification shields SARS-CoV-2 from host antiviral responses such as MDA5 recognition and IFIT1 restriction (Bergant et al., 2022; Russ et al., 2022). Importantly, nsp16 inhibits global host mRNA splicing, which reduces host protein and mRNA levels (Banerjee et al., 2020). Additionally, nsp16 enhances SARS-CoV-2 cell entry by promoting TMPRSS2 expression (Han et al., 2023), highlighting its multifaceted role as a virulence factor. Moreover, nsp16-deficient strains exhibit reduced pathogenicity and induce a strong immune response, suggesting their potential as a live-attenuated vaccine candidate (Ye et al., 2022). Given nsp16’s critical functions, drugs targeting nsp16 or the nsp16-nsp10 complex have been developed (Klima et al., 2022; Nguyen et al., 2022). Therefore, understanding host factors that interact with nsp16 is crucial for uncovering novel therapeutic strategies.

Revised version: Severe acute respiratory syndrome coronavirus 2 (SARS-CoV-2) remains a global public health threat, with over 641 million coronavirus disease 2019 (COVID-19) cases and 6.6 million deaths as of December 2022 (Carabelli et al., 2023; Markov et al., 2023). The virus encodes four structural proteins, 16 non-structural proteins (nsps) involved in replication and transcription, and several accessory proteins (ORFs) linked to immune evasion (V’kovski et al., 2021; Wang et al., 2020). Many viruses, including Ebola (Valle et al., 2021), dengue (Jung et al., 2018), and reovirus (Furuichi et al., 1975), encode 2’-O methyltransferases (2’-O-MTases) that mimic the host 5’ cap structure to evade innate immune recognition (Daffis et al., 2010; Züst et al., 2011). Similarly, SARS-CoV-2 nsp16, in complex with nsp10, functions as a 2′-O-MTase, modifying capped viral RNA from ‘cap-0’ to ‘cap-1’ (Benoni et al., 2021; Park et al., 2022). This modification protects the virus from host antiviral defences, such as MDA5 recognition and IFIT1 restriction (Bergant et al., 2022; Russ et al., 2022). Beyond immune evasion, nsp16 disrupts global host mRNA splicing, reducing cellular protein and mRNA levels (Banerjee et al., 2020). Nsp16 also enhances SARS-CoV-2 entry by upregulating TMPRSS2 expression (Han et al., 2023), underscoring its role as a virulence factor. Notably, nsp16-deficient strains exhibit low pathogenicity and elicit robust immune responses, suggesting their potential as live-attenuated vaccines (Ye et al., 2022). Given the critical functions of nsp16, inhibitors targeting nsp16 or the nsp16-nsp10 complex have been developed (Klima et al., 2022; Nguyen et al., 2022). Therefore, understanding host factors that regulate nsp16 is essential for identifying new therapeutic strategies.

Earlier version: The ubiquitin-proteasome system (UPS) is an important pathway for the targeted degradation of intracellular proteins and is referred to as “the molecular kiss of death” (Dongdem & Wezena, 2021; Park et al., 2020). Substrate ubiquitination requires a cascade of three ligases: E1, E2, and E3 (Zheng & Shabek, 2017). Some E3 ubiquitin ligases function as monomers, while others form multi-subunit E3 complexes. Regardless of their form, all E3 ligases share the same functional specificity in recognizing substrates and determining the specificity of the UPS. In some complexes, an additional protein may be required as a substrate receptor to recognize the substrate (Iconomou & Saunders, 2016; Yang et al., 2021). E3 ligases are indispensable for direct binding to substrates and determining the specificity of the UPS (Li et al., 2021; Wang et al., 2022). Researchers have identified over 600 human genome-encoded E3 ligases that operate in diverse cellular environments, respond to a wide range of cellular signals, and regulate a vast array of protein substrates (Dongdem & Wezena, 2021; Garcia-Barcena et al., 2020). Dysregulation of E3 ligase activity has been implicated in various human diseases, making them promising drug targets (Humphreys et al., 2021). E3 ligases are classified into three classes based on conserved domains and ubiquitin (Ub) transfer mechanisms: Really Interesting New Gene (RING) E3s, homologous to the E6AP carboxyl terminus (HECT) E3s, and RING-between-RING (RBR) E3s (Garcia-Barcena et al., 2020; Morreale & Walden, 2016). RING E3 ligases directly transfer Ub from the E2-Ub complex to the substrate (Deshaies & Joazeiro, 2009), whereas HECT-type E3s transfer Ub to their own catalytic cystine residue before attaching it to the substrate (Huibregtse et al., 1995). RBR E3s combine the characteristics of both RING and HECT types (Walden & Rittinger, 2018). RBR and RING E3s share the RING-binding domains, but RBR family members can generate thioester intermediates with Ub, as is the case with HECT-type E3s (Garcia-Barcena et al., 2020).

Revised version: The ubiquitin-proteasome system (UPS) is a key pathway for targeted protein degradation, often described as ‘the molecular kiss of death’ (Dongdem & Wezena, 2021; Park et al., 2020). Substrate ubiquitination involves a cascade of three enzymes: E1 activating enzyme, E2 conjugating enzyme, and E3 ubiquitin ligase (Zheng & Shabek, 2017). Some E3 ubiquitin ligases function as monomers, whereas others assemble into multi-subunit complexes. Regardless of their structure, all E3 ligases play a crucial role in substrate recognition, dictating UPS specificity. In certain complexes, an additional protein may function as a substrate receptor, mediating substrate-specific recognition (Iconomou & Saunders, 2016; Yang et al., 2021; Li et al., 2021; Wang et al., 2022). Over 600 E3 ligases encoded in the human genome regulate diverse cellular processes in response to various signals (Dongdem & Wezena, 2021; Garcia-Barcena et al., 2020). Dysregulation of their activity is linked to multiple diseases, making them attractive therapeutic targets (Humphreys et al., 2021). Based on conserved domains and ubiquitin (Ub) transfer mechanisms, E3 ligases are categorised into three classes: Really Interesting New Gene (RING), homologous to the E6AP carboxyl terminus (HECT), and RING-between-RING (RBR) ligases (Garcia-Barcena et al., 2020; Morreale & Walden, 2016). RING E3s transfer Ub directly from the E2-Ub complex to substrates (Deshaies & Joazeiro, 2009), whereas HECT-type E3s first attach Ub to their own catalytic cysteine before transferring it (Huibregtse et al., 1995). These distinct mechanisms highlight the complexity and versatility of E3 ligase function in cellular regulation (Walden & Rittinger, 2018; Garcia-Barcena et al., 2020).

Earlier version: To investigate the relationship between SARS-CoV-2 and host factors, we focused on the interaction between nsp16 and the UPS. In this study, we present a novel regulatory mechanism targeting nsp16 via the host UPS. We demonstrate, for the first time, that nsp16 is ubiquitinated and degraded by the host proteasome. Interestingly, two E3 ligases from distinct families, the RING-type MARCHF7 and the HECT-type UBR5, independently recognise and target nsp16 for degradation in the cytoplasm and nucleus. Since nsp16 possesses 2’-O-MTase activity, its enzymatic function converts the viral RNA 5’ cap structure from cap-0 to cap-1. This dual targeting of nsp16 by MARCHF7 and UBR5 leads to its degradation, disrupting its function and ultimately demonstrating effective antiviral activity against SARS-CoV-2. Our findings identify novel therapeutic targets for developing effective strategies to prevent SARS-CoV-2 and treat COVID-19.

Revised version: To investigate the relationship between SARS-CoV-2 and host factors, we focused on UPS proteins that interact with nsp16. In this study, we demonstrate for the first time that nsp16 undergoes ubiquitination and proteasomal degradation. Notably, two E3 ligases from distinct families—RING-type MARCHF7 and HECT-type UBR5—independently target nsp16 for degradation, disrupting its function and exhibiting strong antiviral activity against SARS-CoV-2. Our findings reveal novel therapeutic targets for combating SARS-CoV-2 and treating COVID-19.

Results

Added in the revised version: SARS-CoV-2 nsp16 is degraded via the ubiquitin-proteasome pathway

Removed in the revised version: Ubiquitination and proteasomal degradation of SARS-CoV-2 nsp16 by the UPS

Earlier version: Ubiquitination plays an important role in SARS-CoV-2 infection and pathogenesis (Gao et al., 2022; Guo et al., 2021; Li et al., 2023; Zhang et al., 2021; Zhang et al., 2024; Zhang et al., 2023). To explore whether the non-structural proteins of SARS-CoV-2 are regulated by the UPS, we examined the effect of the proteasome inhibitor MG132 on the protein expression of all 16 non-structural proteins. MG132 treatment markedly increased the abundance of nsp8, nsp11, and nsp16 proteins in HEK293T cells (Fig. 1A), suggesting their potential degradation by the UPS. It has been reported that nsp8 undergoes UPS-mediated degradation via TRIM22-mediated ubiquitination (Fan et al., 2024). In this study, we focused on nsp16 for further investigation. To confirm the role of the UPS in nsp16 degradation, we employed additional proteasome inhibitors, Bortezomib and Carfilzomib. Unlike the lysosomal inhibitors Bafilomycin A1 and NH4Cl and the autophagy-lysosomal inhibitor vinblastine, these proteasome inhibitors enhanced nsp16 stability (Fig. 1B). The impact of the USP on the half-life of nsp16 was investigated. To this end, nsp16-expressing cells were treated with a protein synthesis inhibitor, cycloheximide, and the kinetics of nsp16 decay were examined. We observed that the half-life of nsp16 in MG132-treated samples was significantly longer than that in MG132-untreated samples (15 h versus 2 h) (Fig. 1C and 1D).

Revised version: Ubiquitination plays a crucial role in SARS-CoV-2 infection and pathogenesis (Gao et al., 2022; Guo et al., 2021; Li et al., 2023; Zhang et al., 2021; Zhang et al., 2024; Zhang et al., 2023). To determine whether SARS-CoV-2 non-structural proteins (nsps) are UPS-regulated, we examined the effect of the proteasome inhibitor, MG132, on nsps expression. MG132 treatment significantly increased the nsp8, nsp11, and nsp16 abundance in HEK293T cells (Fig. 1A), suggesting UPS-mediated degradation. Nsp8 undergoes TRIM22-mediated ubiquitination and proteasomal degradation (Fan et al., 2024). In this study, we focused on nsp16 for further investigation. To confirm UPS-mediated nsp16 degradation, we tested additional proteasome inhibitors, Bortezomib and Carfilzomib, both of which enhanced nsp16 stability. In contrast, lysosomal inhibitors (Bafilomycin A1 and NH4Cl) and the autophagy-lysosomal inhibitor, Vinblastine, had no such effect (Fig. 1B). To further assess the role of the UPS in regulating nsp16 stability, we treated nsp16-expressing cells with cycloheximide, a protein synthesis inhibitor, and measured the half-life of nsp16. MG132 treatment significantly prolonged nsp16 stability, extending its half-life from 2 h to 15 h (Fig. 1C and 1D).

The non-structural protein nsp16 of SARS-CoV-2 was identified that can be degraded through the proteasome pathway.

Earlier version: A. The non-structural proteins nsp8, nsp11 and nsp16 could be restored by the proteasome inhibitor MG132. HEK293T cells in 12-well plates were transfected with the plasmids of 16 nonstructural proteins (nsp1-16) encoded by SARS-CoV-2. Thirty-six hours later, the cells were treated with MG132 (10 µM) or DMSO for 12 h before collection. The protein level was detected by Immunoblotting (IB). Quantification of nsp protein levels relative to the control protein is shown. Data are representative of three independent experiments and shown as average ±SD (n = 3). Significance was determined by a two-tailed t-test: *P < 0.05; **P < 0.01; ***P < 0.001. B. Proteasomal inhibitors but no other inhibitors stabilized nsp16 protein. HEK293T cells transfected with the nsp16-Flag expression vector were treated with dimethyl sulfoxide (DMSO), MG132 (10 µM), Bortezomib (10 µM), Carfilzomib (10 µM), Bafilomycin A1 (5 µM), Vinblastine (2.5 µM), or NH4CL (2.5 µM) for 12 h prior to harvest. The cell lysates were analyzed by anti-Flag antibody. (C-D). The half-life of nsp16 was prolonged by the proteasome inhibitor MG132. C. HEK293T cells were transfected with the nsp16-Flag-expressing plasmids. 12 hours later, the cells were treated with DMSO or MG132 (10 µM) for 12 h, then 50 µg/mL cycloheximide (CHX) was added. Cells were harvested at the indicated times to detect the level of viral protein by anti-Flag antibody. D. Quantification of nsp16 protein levels relative to tubulin at different time points is shown. The half-life of the nsp16 protein was determined based on protein quantification using Image J, combined with the protein half-life formula for calculation. Results are shown as mean ± SD (n = 3 independent experiments). ***, P < 0.001 by by a two-tailed t-test.

Revised version: A. The non-structural proteins nsp8, nsp11 and nsp16 could be restored by the proteasome inhibitor MG132. HEK293T cells in 12-well plates were transfected with the plasmids of 16 nonstructural proteins (nsp1-16) encoded by SARS-CoV-2. Thirty-six hours later, the cells were treated with MG132 (10 µM) or DMSO for 12 h before collection. The protein level was detected by Immunoblotting (IB). Quantification of nsp protein levels relative to the control protein is shown. Data are representative of three independent experiments and shown as average ±SD (n = 3). Significance was determined by a two-tailed t-test: *P < 0.05; **P < 0.01; ***P < 0.001. B. Proteasomal inhibitors but no other inhibitors stabilized nsp16 protein. HEK293T cells transfected with the nsp16-Flag expression vector were treated with dimethyl sulfoxide (DMSO), MG132 (10 µM), Bortezomib (10 µM), Carfilzomib (10 µM), Bafilomycin A1 (5 µM), Vinblastine (2.5 µM), or NH4CL (2.5 µM) for 12 h prior to harvest. The cell lysates were analyzed by anti-Flag antibody. (C-D). The half-life of nsp16 was prolonged by the proteasome inhibitor MG132. C. HEK293T cells were transfected with the nsp16-Flag-expressing plasmids. 12 hours later, the cells were treated with DMSO or MG132 (10 µM) for 12 h, then 50 µg/mL cycloheximide (CHX) was added. Cells were harvested at the indicated times to detect the level of viral protein by anti-Flag antibody. D. Quantification of nsp16 protein levels relative to tubulin at different time points is shown. The half-life of the nsp16 protein was determined based on protein quantification using Image J, combined with the protein half-life formula for calculation. Results are shown as mean ± SD (n = 3 independent experiments). ***, P < 0.001 by by a two-tailed t-test. E. Samples were prepared for mass spectrometry, and nsp16 interacting proteins were obtained by immunoprecipitation (IP) (created with BioRender.com and the agreement no. is LC27PCBEOF). The plasmids were transfected into HEK293T cells for 48 h. Treat cells with or without MG132 (10 µM) for 12 h prior to harvest. The whole-cell lysates were incubated with protein G agarose beads conjugated with anti-Flag antibodies and used for IB with anti-Flag antibodies to detect the nsp16 protein. Samples enriched for proteins were analyzed by mass spectrometry.

Removed in the revised version: E. Samples were prepared for mass spectrometry, and nsp16 interacting proteins were obtained by immunoprecipitation (IP) (created using BioRender.com). The plasmids were transfected into HEK293T cells for 48 h. Treat cells with or without MG132 (10 µM) for 12 h prior to harvest. The whole-cell lysates were incubated with protein G agarose beads conjugated with anti-Flag antibodies and used for IB with anti-Flag antibodies to detect the nsp16 protein. Samples enriched for proteins were analyzed by mass spectrometry.

Earlier version: To identify proteins interacting with nsp16, we performed co-immunoprecipitation (Co-IP) followed by mass spectrometry (MS) analysis, comparing nsp16-expressing cells treated with and without MG132 (Fig. 1E). KEGG pathway and Gene Ontology analyses of the MS data (Figure 1—figure supplement 1A). Biological process enrichment analysis showed that blocking nsp16 degradation resulted in an increase in the number of proteins that interact with it, including 43 proteins involved in virus-associated processes and 22 proteins involved in regulating mRNA stability. Five proteins were related to the regulation of the cellular antiviral defence. Additionally, interactions were associated with mRNA splicing and RNA methylation (Banerjee et al., 2020), consistent with nsp16’s role as a 2’-O-MTase. Furthermore, SARS-CoV-2 uses nsp16 to disrupt host mRNA splicing to facilitate its infection (Park et al., 2022; Russ et al., 2022; Züst et al., 2011). These findings corroborate the reliability of our MS data.

Revised version: To identify nsp16-interacting proteins, we performed co-immunoprecipitation (Co-IP) followed by mass spectrometry (MS) analysis, comparing nsp16-expressing cells with and without MG132 treatment (Fig. 1E). Kyoto Encyclopedia of Genes and Genomes (KEGG) pathway and Gene Ontology (GO) analyses (Figure 1—figure supplement 1A) revealed that inhibiting nsp16 degradation increased its interacting proteins, including 43 associated with viral processes, 22 involved in mRNA stability regulation, and five linked to cellular antiviral defence. Additionally, interactions related to mRNA splicing and RNA methylation were identified (Banerjee et al., 2020), consistent with the function of nsp16 as a 2′-O-MTase. Furthermore, SARS-CoV-2 exploits nsp16 to disrupt host mRNA splicing, which promotes infection (Park et al., 2022; Russ et al., 2022; Züst et al., 2011), further validating our MS data.

Earlier version: As expected, we found that several nsp16-binding proteins were related to ubiquitination and degradation pathways (Figure 1 — figure supplement 1B). Four deubiquitinases (DUBs), the E2 ligase UBE2D3, and 14 proteasomal enzymes were identified. Of these, six E3 ligases, UBR5, MARCHF7, HECTD1, TRIM32, MYCBP2, and TRIM21, were chosen for further investigation.

Revised version: As expected, several nsp16-binding proteins were associated with ubiquitination and degradation pathways (Figure 1-figure supplement 1B). We identified four deubiquitinases (DUBs), the E2 ligase UBE2D3, and 14 proteasomal enzymes. Among them, six E3 ligases—UBR5, MARCHF7, HECTD1, TRIM32, MYCBP2, and TRIM21—were selected for further investigation.

E3 ligases, UBR5 and MARCHF7, independently mediate nsp16 degradation

Earlier version: To identify the E3 ligases responsible for nsp16 degradation, we designed small interfering RNA (siRNA) targeting the six candidate E3 ligases and evaluated their effects on nsp16 abundance. Knockdown of UBR5 and MARCHF7 stabilised the protein levels of nsp16 (Fig. 2A). Furthermore, we generated stable cell lines with UBR5 or MARCHF7 knockdowns to further investigate their effects on nsp16 stability and confirm silencing efficiency (Fig. 2B). To determine whether UBR5 and MARCHF7 cooperate in nsp16 degradation, we knocked down UBR5 or MARCHF7 using siRNA in MARCHF7- or UBR5-knockdown cells, respectively. Knockdown of UBR5 in MARCHF7-knockdown cells, and vice versa, further enhanced nsp16 stability. This indicates that UBR5 and MARCHF7 act independently to degrade nsp16 (Fig. 2C). We further validated the independent function of these E3 ligases by examining nsp16 degradation. MARCHF7 overexpression induced nsp16 degradation in both MARCHF7-knockdown and MARCHF7/UBR5-double-knockdown cells. Similarly, UBR5 overexpression degraded nsp16 even in the presence of MARCHF7 knockdown. Knockdown efficiencies were confirmed for all experiments (Figure 2 — figure supplement 1A and 1B). Notably, overexpression of wild-type MARCHF7 or UBR5, but not mutants lacking the functional RING domain (1-542 aa) or the HECT domain inactivated mutant, disrupts the stability of nsp16 in cells with MARCHF7 or UBR5 knockdown (Fig. 2D and 2E). These findings demonstrate that UBR5 and MARCHF7 independently ubiquitinate nsp16 via their HECT or RING domains, targeting it for proteasomal degradation.

Revised version: To identify the E3 ligases responsible for nsp16 degradation, we designed small interfering RNAs (siRNAs) targeting six candidate E3 ligases and assessed their effects on nsp16 abundance. UBR5 and MARCHF7 knockdown significantly stabilised nsp16 protein levels (Fig. 2A). To further investigate their roles in nsp16 degradation, we generated stable cell lines with UBR5 or MARCHF7 knockdowns and confirmed silencing efficiency (Fig. 2B). To determine whether UBR5 and MARCHF7 function cooperatively, we performed dual knockdowns. Silencing UBR5 in MARCHF7-knockdown cells, and vice versa, further enhanced nsp16 stability, indicating that these ligases act independently (Fig. 2C). Overexpression studies further supported their independent roles—MARCHF7 overexpression induced nsp16 degradation even in MARCHF7-knockdown and MARCHF7/UBR5 double-knockdown cells. Similarly, UBR5 overexpression promoted nsp16 degradation despite MARCHF7 knockdown. Silencing efficiencies were validated across all experiments (Figure 2 — figure supplement 1A and 1B). Notably, wild-type MARCHF7 and UBR5, but not their mutants lacking functional RING (1-542 aa) or HECT domains, respectively, effectively degraded nsp16 in knockdown cells (Fig. 2D and 2E). These findings confirm that UBR5 and MARCHF7 independently mediate nsp16 ubiquitination via their HECT and RING domains, targeting it for proteasomal degradation.

MARCHF7 and UBR5 were identified as E3 ubiquitin ligases involves in nsp16 protein degradation.

Earlier version: A. Knockdown of MARCHF7 or UBR5 resulted in nsp16 restoration. HEK293T cells were transfected with siRNA of E3 ligase candidates for 24 h, followed by co-incubation with the nsp16-Flag-expressing plasmids for 48 h, treated with MG132 (10 µM) for 16h before harvesting, lysed, and subjected to IB assay using anti-Flag antibody. RT-qPCR was conducted to determine the mRNA expression levels of E3 ligase candidates. The si-RNA targeting regions for the candidate E3 ubiquitin ligase proteins and the targeted regions for RT-qPCR are shown in Appendix-figure S1A. Data are representative of three independent experiments and shown as average ± SD (n = 3). Significance was determined by a two-tailed t-test: ***P < 0.001.

Revised version: A. Knockdown of MARCHF7 or UBR5 resulted in nsp16 restoration. HEK293T cells were transfected with siRNA of E3 ligase candidates for 24 h, followed by co-incubation with the nsp16-Flag-expressing plasmids for 48 h, treated with MG132 (10 µM) for 16h before harvesting, lysed, and subjected to IB assay using anti-Flag antibody. RT-qPCR was conducted to determine the mRNA expression levels of E3 ligase candidates. The si-RNA targeting regions for the candidate E3 ubiquitin ligase proteins and the targeted regions for RT-qPCR are shown in Appendix-figure S1A. Data are representative of three independent experiments and shown as average ± SD (n = 3). Significance was determined by a two-tailed t-test: ***P < 0.001. B. RNA levels of UBR5 or MARCHF7 from HEK293T cells infected with lentivirus containing control or shRNA targeting UBR5 or MARCHF7 for 48 h and screened with antibiotics for 48 h. Knockdown cell lines were transfected with plasmids expressing nsp16-Flag, collected at the indicated times, and the protein levels of nsp16, MARCHF7, and UBR5 were detected by IB.C. MARCHF7 and UBR5 acted separately and did not depend on each other. HEK293T cells stably expressing UBR5 shRNA or MARCHF7 shRNA were transfected with siRNA of MARCHF7 or UBR5 for 24 h, respectively, followed by co-incubation with the nsp16-Flag-expressing plasmids for 48 h. The protein levels and the RNA levels of nsp16, UBR5 and MARCHF7 were measured by IB and RT-qPCR, respectively. (D-E). In HEK293T cells stably expressing UBR5 shRNA or MARCHF7 shRNA, nsp16 was degraded by overexpressed UBR5 or MARCHF7, respectively, whereas the mutant failed to degrade nsp16. The cell lysates were analyzed by anti-Flag antibody.

Removed in the revised version: B. RNA levels of UBR5 or MARCHF7 from HEK293T cells infected with lentivirus containing control or shRNA targeting UBR5 or MARCHF7 for 48 h and screened with antibiotics for 48 h. Knockdown cell lines were transfected with plasmids expressing nsp16-Flag, collected at the indicated times, and the protein levels of nsp16, MARCHF7, and UBR5 were detected by IB.C. MARCHF7 and UBR5 acted separately and did not depend on each other. HEK293T cells stably expressing UBR5 shRNA or MARCHF7 shRNA were transfected with siRNA of MARCHF7 or UBR5 for 24 h, respectively, followed by co-incubation with the nsp16-Flag-expressing plasmids for 48 h. The protein levels and the RNA levels of nsp16, UBR5 and MARCHF7 were measured by IB and RT-qPCR, respectively.

Removed in the revised version: (D-E). In HEK293T cells stably expressing UBR5 shRNA or MARCHF7 shRNA, nsp16 was degraded by overexpressed UBR5 or MARCHF7, respectively, whereas the mutant failed to degrade nsp16. The cell lysates were analyzed by anti-Flag antibody.

UBR5 and MARCHF7 mediate distinct Ub linkages on nsp16

Earlier version: We first confirmed the ubiquitination of nsp16, both in the presence and absence of exogenous Ub (Fig. 3A-B). To further investigate the regulation by E3 ligases, we examined the impact of UBR5 or MARCH7 knockdown on nsp16 ubiquitination. Knockdown of either UBR5 or MARCHF7 decreased the ubiquitination level of nsp16 compared to that in the negative control (Fig. 3C). Ub contains seven lysine residues (K6, K11, K27, K29, K33, K49, and K63) that form polyubiquitin chains with distinct functions, ultimately dictating the fate of the attached target protein (Grice & Nathan, 2016). To determine the specific type of polyubiquitin chain modification of nsp16 mediated by UBR5 and MARCH7, we constructed a series of Ub mutants retaining only a single lysine residue. Except for K33, all single-lysine Ub mutants promoted nsp16 ubiquitylation to varying extents, suggesting a complex polyubiquitin chain structure on nsp16 potentially regulated by multiple E3 ligases or E2-E3 ligase pairs (Fig. 3D). Notably, analysis of specific lysine linkages revealed that MARCHF7 knockdown primarily reduced the K27-linked ubiquitination of nsp16, while UBR5 knockdown primarily reduced the K48-linked ubiquitination of nsp16 (Fig. 3E and 3F). These findings indicate that MARCHF7 and UBR5 induce K27- and K48-linked ubiquitination on nsp16, respectively.

Revised version: We first confirmed nsp16 ubiquitination both with and without exogenous ubiquitin (Ub) expression (Fig. 3A-B). To investigate the role of E3 ligases in this process, we examined the impact of UBR5 and MARCHF7 knockdown on nsp16 ubiquitination. Silencing either ligase reduced nsp16 ubiquitination compared to the negative control, indicating their involvement in this modification (Fig. 3C). Ub contains seven lysine residues (K6, K11, K27, K29, K33, K48, and K63), which form polyubiquitin chains with distinct functions that determine the fate of the attached protein (Grice & Nathan, 2016). To identify the specific polyubiquitin chain types formed on nsp16 by UBR5 and MARCHF7, we used a series of Ub mutants, each retaining only a single lysine residue. Except for K33, all single-lysine Ub mutants supported nsp16 ubiquitination to varying degrees, suggesting a complex ubiquitin architecture potentially regulated by multiple E3 ligases or E2-E3 pairs (Fig. 3D). Further analysis of lysine-specific linkages revealed that MARCHF7 primarily mediates K27-linked ubiquitination of nsp16, whereas UBR5 facilitates K48-linked ubiquitination (Fig. 3E and 3F). These findings establish that UBR5 and MARCHF7 independently regulate nsp16 degradation through distinct ubiquitin linkages.

MARCHF7 or UBR5 catalyze the formation of K-27 type or K-48 type ubiquitin chains of nsp16 respectively.

Removed in the revised version: A. Nsp16 can be ubiquitinated. HEK293T cells co-transfected with ubiquitin-Myc and nsp16-Flag or transfected with nsp16-Flag alone. The cells were treated with MG132 for 12 h before collection. The whole-cell lysates were incubated with anti-Flag beads and used for IB with anti-Myc or anti-Flag antibodies to detect the polyubiquitination chain of nsp16.

Removed in the revised version: B. Assess the endogenous ubiquitination level of nsp16 protein. Cells were transfected with nsp16-Flag or an empty vector, and collected 48 hours later. Prior to harvesting, cells were treated with MG132 for 16 hours. Co-IP experiments were then performed to analyze the endogenous ubiquitination level of nsp16.

Earlier version: C. The level of ubiquitination of nsp16 decreased with decreasing the protein levels of MARCHF7 or UBR5. E3 was knocked down by transfection with siRNA targeting UBR5 or MARCHF7, and 24 h later ubiquitin-Myc and nsp16-HA were co-transfected or nsp16-HA alone. Cells were treated with MG132 for 16 h before collection. Whole cell lysates were incubated with anti-HA beads, and polyubiquitinated chains of nsp16 were detected by IB with anti-Myc or anti-HA antibodies.

Revised version: A. Nsp16 can be ubiquitinated. HEK293T cells co-transfected with ubiquitin-Myc and nsp16-Flag or transfected with nsp16-Flag alone. The cells were treated with MG132 for 12 h before collection. The whole-cell lysates were incubated with anti-Flag beads and used for IB with anti-Myc or anti-Flag antibodies to detect the polyubiquitination chain of nsp16. B. Assess the endogenous ubiquitination level of nsp16 protein. Cells were transfected with nsp16-Flag or an empty vector, and collected 48 hours later. Prior to harvesting, cells were treated with MG132 for 16 hours. Co-IP experiments were then performed to analyze the endogenous ubiquitination level of nsp16. C. The level of ubiquitination of nsp16 decreased with decreasing the protein levels of MARCHF7 or UBR5. E3 was knocked down by transfection with siRNA targeting UBR5 or MARCHF7, and 24 h later ubiquitin-Myc and nsp16-HA were co-transfected or nsp16-HA alone. Cells were treated with MG132 for 16 h before collection. Whole cell lysates were incubated with anti-HA beads, and polyubiquitinated chains of nsp16 were detected by IB with anti-Myc or anti-HA antibodies. D. Nsp16 can be modified by a variety of ubiquitin chains. HEK293T cells were transfected with either nsp16-HA alone or together with plasmids encoding various mutants of ubiquitin (K6 only, K11 only, K27 only, K29 only, K33 only, K48 only, K63 only). Thirty-six hours later, cells were treated with MG132 for 12 h. Cell lysates were then subjected to immunoprecipitation, followed by IB to analysis. (E-F) MARCHF7 or UBR5 causes nsp16 to be modified by the K27 type or K48 type ubiquitin chain. 293T cell lines with or without MARCHF7 or UBR5 knockdown were co-transfected with plasmids encoding ubiquitin-WT or various mutants of ubiquitin (K6 only, K11 only, K27 only, K29 only, K33 only, K48 only, K63 only). The other experimental methods were the same as C.

Removed in the revised version: D. Nsp16 can be modified by a variety of ubiquitin chains. HEK293T cells were transfected with either nsp16-HA alone or together with plasmids encoding various mutants of ubiquitin (K6 only, K11 only, K27 only, K29 only, K33 only, K48 only, K63 only). Thirty-six hours later, cells were treated with MG132 for 12 h. Cell lysates were then subjected to immunoprecipitation, followed by IB to analysis.

Removed in the revised version: (E-F) MARCHF7 or UBR5 causes nsp16 to be modified by the K27 type or K48 type ubiquitin chain. 293T cell lines with or without MARCHF7 or UBR5 knockdown were co-transfected with plasmids encoding ubiquitin-WT or various mutants of ubiquitin (K6 only, K11 only, K27 only, K29 only, K33 only, K48 only, K63 only). The other experimental methods were the same as C.

UBR5 and MARCHF7 directly interact and colocalise with nsp16 in the endoplasmic reticulum (ER)

Earlier version: MS analysis suggested interactions between nsp16 and both UBR5 and MARCH7. Co-IP experiments confirmed that both Myc-tagged MARCHF7 and endogenous UBR5 interact with nsp16 (Figure 4 — figure supplement 1A and 1B). To determine whether MARCHF7 or UBR5 directly interact with nsp16, fluorescence resonance energy transfer (FRET) assays were performed. When nsp16-YFP was bleached, the fluorescence signals from CFP-UBR5 or CFP-MARCHF7 fusion proteins increased, indicating a direct interaction between these E3 ligases and nsp16 (Figure 4—figure supplement 1C). Image J software was used to quantify the relative fluorescence intensities (Figure 4 — figure supplement 1D). We next investigated whether UBR5 and MARCHF7 require each other for nps16 binding. We compared the binding ability of UBR5 to nsp16 in the presence or absence of MARCHF7 using Co-IP. MARCHF7 knockdown had no effect on the interaction between UBR5 and nsp16, and vice versa (Fig. 4A and 4B), indicating independent binding of UBR5 and MARCHF7 to nsp16. Previous studies have shown that nsp16 is localised to both the nucleus and cytoplasm (Zhang et al., 2020). UBR5 (containing two nuclear localisation signals) and MARCHF7 are present in both compartments (Muñoz-Escobar et al., 2015; Shearer et al., 2018) (Nathan et al., 2008). Immunofluorescence staining revealed that UBR5 and MARCHF7 colocalised primarily with nsp16 in the cytoplasm, with some colocalisation in the nucleus of Hela cells (Figure 4—figure supplement 1E). Similar results were observed in nsp16-transfected HEK293T cells using antibodies against endogenous UBR5 or MARCHF7 (Figure 4 — figure supplement 1F).

Revised version: MS analysis suggested interactions between nsp16 and both UBR5 and MARCHF7. Co-IP experiments confirmed that Myc-tagged MARCHF7 and endogenous UBR5 bind to nsp16 (Figure 4—figure supplement 1A and 1B). To assess whether these interactions are direct, we performed fluorescence resonance energy transfer (FRET) assays. When nsp16-YFP was bleached, the fluorescence intensity of CFP-UBR5 and CFP-MARCHF7 increased, confirming direct interactions (Figure 4 — figure supplement 1C). ImageJ software was used to quantify relative fluorescence intensities (Figure 4 — figure supplement 1D). To determine whether UBR5 and MARCHF7 depend on each other for nsp16 binding, we examined the effect of MARCHF7 knockdown on UBR5-nsp16 interactions, and vice versa, using Co-IP. The interactions remained unchanged, indicating that UBR5 and MARCHF7 independently bind to nsp16 (Fig. 4A and 4B). Nsp16 localises to both the nucleus and cytoplasm (Zhang et al., 2020). UBR5 (containing two nuclear localization signals) and MARCHF7 are also present in both compartments (Muñoz-Escobar et al., 2015; Shearer et al., 2018) (Nathan et al., 2008). Immunofluorescence staining revealed that UBR5 and MARCHF7 colocalised with nsp16 predominantly in the cytoplasm, with some nuclear colocalization observed in HeLa cells (Figure 4—figure supplement 1E). Similar results were obtained in nsp16-transfected HEK293T cells using antibodies against endogenous UBR5 and MARCHF7 (Figure 4—figure supplement 1F).

MARCHF7 and UBR5 directly interact with nsp16 respectively.

Removed in the revised version: (A-B). The binding of MARCHF7 or UBR5 to nsp16 was not mutually dependent. The binding of nsp16 to UBR5 or MARCHF7 was identified by co-immunoprecipitation in HEK293T cells transfected siMARCHF7 or siUBR5, respectively. The immunoprecipitates and input were analyzed by IB. The knockdown efficiency was detected by RT-qPCR and IB.

Earlier version: (C-D). MARCHF7 or UBR5 co-localized with nsp16 in the endoplasmic reticulum. Hela cells were co-transfected with YFP-nsp16(yellow) and CFP-UBR5(cyan) or CFP-MARCHF7(cyan). The organelles were labeled with antibodies against marker proteins of endoplasmic reticulum, Golgi apparatus and mitochondria respectively(red). The cells were analyzed by confocal microscopy (C). Scale bars, 20 um. The ratio of colocalization was quantified by measuring the fluorescence intensities using Image J (D).

Revised version: (A-B). The binding of MARCHF7 or UBR5 to nsp16 was not mutually dependent. The binding of nsp16 to UBR5 or MARCHF7 was identified by co-immunoprecipitation in HEK293T cells transfected siMARCHF7 or siUBR5, respectively. The immunoprecipitates and input were analyzed by IB. The knockdown efficiency was detected by RT-qPCR and IB. (C-D). MARCHF7 or UBR5 co-localized with nsp16 in the endoplasmic reticulum. Hela cells were co-transfected with YFP-nsp16(yellow) and CFP-UBR5(cyan) or CFP-MARCHF7(cyan). The organelles were labeled with antibodies against marker proteins of endoplasmic reticulum, Golgi apparatus and mitochondria respectively(red). The cells were analyzed by confocal microscopy (C). Scale bars, 20 um. The ratio of colocalization was quantified by measuring the fluorescence intensities using Image J (D).

Earlier version: To identify the specific cellular compartment where UBR5 or MARCH7 interacts with nsp16, we co-transfected cells with UBR5-CFP or MARCH7-CFP along with nsp16-YFP. Immunostaining was performed with organelle-specific antibodies to label the mitochondrial marker COX5A, the ER marker PDI, and the Golgi apparatus marker GM130. Notably, both UBR5 and MARCHF7 interacted with nsp16 and colocalised with PDI, but not with COX5A or GM130 (Fig. 4C and 4D). These findings indicate that UBR5 and MARCHF7 directly interact with nsp16 in the ER.

Revised version: To pinpoint the specific cellular compartment where these interactions occur, we co-transfected UBR5-CFP or MARCHF7-CFP with nsp16-YFP and performed immunostaining using organelle-specific markers: COX5A (mitochondria), PDI (ER), and GM130 (Golgi apparatus). Notably, UBR5 and MARCHF7 both interacted with nsp16 and colocalised with PDI, but not with COX5A or GM130 (Fig. 4C and 4D), indicating that both E3 ligases interact with nsp16 primarily in the ER.

Functional domains of UBR5 and MARCHF7 are required for nsp16 interaction and ubiquitination

Earlier version: UBR5 is a four-domain E3 ligase containing two nuclear localisation signals. The four domains are UBA, UBR, PABC, and HECT (Muñoz-Escobar et al., 2015) (Figure 4—figure supplement 2A). Notably, the HECT domain is essential for its E3 ligase activity, where UBR5 must be conjugated to Ub before transferring it to the substrate (Kim et al., 2021).

Revised version: UBR5 is a four-domain E3 ligase containing two nuclear localization signals. Its domains include UBA, UBR, PABC, and HECT (Muñoz-Escobar et al., 2015) (Figure 4— figure supplement 2A). Notably, the HECT domain is critical for its E3 ligase activity, as UBR5 must first conjugate to Ub before transferring it to substrates (Kim et al., 2021). To determine which UBR5 domain mediates nsp16 ubiquitination, we used UBR5 mutants with individual domain inactivations. Only the HECT domain mutant failed to degrade nsp16 (Figure 4—figure supplement 2B), consistent with previous studies (Zhou et al., 2022). This mutant also lost the ability to ubiquitinate nsp16, confirming the essential role of the HECT domain in this process (Figure 4—figure supplement 2C).

Removed in the revised version: To identify the UBR5 domain responsible for nsp16 ubiquitination, we utilised UBR5 mutants with inactivated individual domains. Only the HECT domain mutant failed to degrade nsp16 (Figure 4—figure supplement 2B), consistent with previous findings (Zhou et al., 2022), and was unable to induce nsp16 ubiquitination, further supporting its essential function (Figure 4—figure supplement 2C).

Earlier version: To identify the MARCHF7 region responsible for nsp16 degradation, we constructed a series of truncation mutants, as previously described (Figure 4 — figure supplement 2D) (Zhao et al., 2018) (Nathan et al., 2008). All mutants lost the ability to degrade nsp16 (Figure 4—figure supplement 2E). Furthermore, we explored the ability of these mutants to interact with nsp16 and their impact on nsp16 ubiquitination. Only the mutant with an intact N-terminal region (aa 1–542) retained strong binding to nsp16, while mutants with the active RING domain region (aa 543–616) did not. Consistently, only the wild-type MARCHF7 induced the formation of K27-linked Ub chains on nsp16 (Figure 4 — figure supplement 2F). These findings suggest that the N-terminal of MARCH7 mediates binding to nsp16, while the RING domain is required for its K27-linked ubiquitination activity.

Revised version: To identify the MARCHF7 region responsible for nsp16 degradation, we generated a series of truncation mutants as described previously (Figure 4—figure supplement 2D) (Zhao et al., 2018) (Nathan et al., 2008). All mutants lost the ability to degrade nsp16 (Figure 4 — figure supplement 2E). Further analysis revealed that only the mutant containing the intact N-terminal region (aa 1–542) strongly bound to nsp16, while those retaining the active RING domain region (aa 543–616) did not. Consistently, only wild-type MARCHF7 promoted K27-linked ubiquitination of nsp16 (Figure 4—figure supplement 2F).

Added in the revised version: These results indicate that the N-terminal region of MARCHF7 is essential for nsp16 binding, whereas its RING domain is required for K27-linked ubiquitination activity.

Earlier version: UBR5 and MARCHF7 mediate antiviral activity against SARS-CoV-2 by targeting nsp16 for degradation

Revised version: UBR5 and MARCHF7 suppress SARS-CoV-2 replication by targeting nsp16 for degradation

Earlier version: To examine the functional impact of nsp16 degradation by UBR5 and MARCHF7 on viral replication, we used a biosafety level 2 cell culture system to generate SARS-CoV-2 transmissible virus-like particles capable of infecting and replicating in Caco2-Nint cells (Ju et al., 2021). Knockdown of either UBR5 or MARCHF7 in these cells significantly increased SARS-CoV-2 replication compared with that in the control group. Notably, simultaneous knockdown of both UBR5 and MARCHF7 further increased SARS-CoV-2 replication (Figure 5—figure supplement 1A and 1B).

Revised version: To assess the impact of nsp16 degradation on viral replication, we used a biosafety level 2 system to generate SARS-CoV-2 transmissible virus-like particles capable of infecting Caco2-Nint cells (Ju et al., 2021). Knockdown of either UBR5 or MARCHF7 significantly increased SARS-CoV-2 replication, with simultaneous knockdown further enhancing viral replication (Figure 5—figure supplement 1A and 1B).

Added in the revised version: To investigate nsp16 modifications following infection, we infected HEK293T-ACE2 cells with the Wuhan strain for 48 h. Co-IP using an anti-nsp16 antibody confirmed its interaction with endogenous MARCHF7 and UBR5, along with ubiquitination (Fig. 5A). In a biosafety level 3 facility, we further validated the antiviral roles of UBR5 and MARCHF7 using the Wuhan (MOI 0.01) and Omicron strains (MOI 0.001). Knockdown of either E3 ligase in Caco2 cells significantly increased intracellular and secreted viral mRNA levels (M and E genes) for both strains. Viral titres in the supernatants also increased, with UBR5 knockdown causing over a 1-log rise for the Wuhan strain and even greater effects for Omicron (Fig. 5C–E and 5F–H for the Wuhan and Omicron strains, respectively). Immunoblotting confirmed increased intracellular and secreted N protein levels upon UBR5 or MARCHF7 knockdown (Fig. 5I), with knockdown efficiencies verified (Fig. 5B). Due to low transfection efficiency in Caco2 cells, we overexpressed UBR5-Myc or MARCH7-Myc in HEK293T-ACE2 cells. Overexpression significantly reduced viral mRNA (M and E genes) and N protein levels in both strains, accompanied by a 0.5-log decrease in viral titres. However, co-transfection with increasing nsp16 levels counteracted these inhibitory effects (Fig. 6A–H and Figure 6—figure supplement 1A–H). To confirm the role of UBR5 and MARCHF7 enzymatic activity, we overexpressed catalytically inactive mutants (UBR5-Δ HECT or MARCHF7-aa 1-542) alongside nsp16. These mutants failed to suppress viral replication. Additionally, gradual increases in nsp16 expression did not further enhance viral replication, with only a slight increase in M mRNA, while E and N protein levels remained unchanged (Figure 6—figure supplement 2A–H). These findings demonstrate that UBR5 and MARCHF7 exert antiviral effects by ubiquitinating and degrading nsp16, thereby significantly inhibiting SARS-CoV-2 replication.

Knockdown of MARCHF7 or UBR5 promotes viral replication.

Removed in the revised version: A. The virus-encoded nsp16 protein interacts with endogenous MARCHF7 and UBR5 and undergoes ubiquitination modification. In 293T-ACE2 cells, with or without Wuhan strain infection (MOI: 0.01), the medium was changed 2 hours post-infection, and cells were harvested 48 hours later, with MG132 treatment added 16 hours before harvesting. nsp16 protein was enriched using Protein-G beads coupled with the nsp16 antibody, and interactions and ubiquitination were analyzed by immunoblotting (IB) with endogenous antibodies against MARCHF7, UBR5, and ubiquitination.

Earlier version: (B-I). MARCHF7 and UBR5 were knocked down by siRNA in Caco2 cells. 24 h after transfection, the cells were infected with Wuhan strain (MOI:0.01) (C-E) or Omicron BA.1 strain (MOI: 0.001) (F-H), respectively. 2 h post infection, the supernatant was discarded, and the cells were cultured in DMEM containing 3% fetal bovine serum for 48 h. The mRNA levels of SARS-CoV-2 M and E genes in the cells (C, F) and E genes in supernatant (D, G) were detected by RT-qPCR and the viral titers in supernatant (E, H) were measured. The N protein levels of Wuhan or Omicron viruses were detected by IB (I). Knock-down efficiencies of MARCHF7 and UBR5 were detected by RT-qPCR or IB (B, I). Data are representative of three independent experiments and shown as average ±SD (n = 3). Significance was determined by one-way ANOVA, followed by a Tukey multiple comparisons posttest: *P < 0.05; **P < 0.01; ***P < 0.001.

Revised version: A. The virus-encoded nsp16 protein interacts with endogenous MARCHF7 and UBR5 and undergoes ubiquitination modification. In 293T-ACE2 cells, with or without Wuhan strain infection (MOI: 0.01), the medium was changed 2 hours post-infection, and cells were harvested 48 hours later, with MG132 treatment added 16 hours before harvesting. nsp16 protein was enriched using Protein-G beads coupled with the nsp16 antibody, and interactions and ubiquitination were analyzed by immunoblotting (IB) with endogenous antibodies against MARCHF7, UBR5, and ubiquitination. (B-I). MARCHF7 and UBR5 were knocked down by siRNA in Caco2 cells. 24 h after transfection, the cells were infected with Wuhan strain (MOI:0.01) (C-E) or Omicron BA.1 strain (MOI: 0.001) (F-H), respectively. 2 h post infection, the supernatant was discarded, and the cells were cultured in DMEM containing 3% fetal bovine serum for 48 h. The mRNA levels of SARS-CoV-2 M and E genes in the cells (C, F) and E genes in supernatant (D, G) were detected by RT-qPCR and the viral titers in supernatant (E, H) were measured. The N protein levels of Wuhan or Omicron viruses were detected by IB (I). Knock-down efficiencies of MARCHF7 and UBR5 were detected by RT-qPCR or IB (B, I). Data are representative of three independent experiments and shown as average ±SD (n = 3). Significance was determined by one-way ANOVA, followed by a Tukey multiple comparisons posttest: *P < 0.05; **P < 0.01; ***P < 0.001.

Removed in the revised version: Given the critical antiviral roles of UBR5 and MARCH7, we investigated the modification of nsp16 following viral infection. HEK293T cells stably expressing the SARS-CoV-2 entry receptor ACE2 (HEK293T-ACE2) were infected with the Wuhan strain for 48 h. Co-IP using an anti-nsp16 antibody revealed that nsp16 interacts with endogenous MARCHF7 and UBR5 and undergoes ubiquitination (Fig. 5A). We further validated the antiviral effects of UBR5 and MARCH7 in a biosafety level 3 facility using the Wuhan strain at a multiplicity of infection (MOI) of 0.01 and the Omicron strain at an MOI of 0.001. Compared to the control group, UBR5 or MARCH7 knockdown in Caco2 cells significantly increased intracellular and secreted viral mRNA levels (M and E genes) of both strains. Additionally, with the Wuhan strain, MARCH7 knockdown resulted in a 0.6-log increase in viral titre in the supernatants, while UBR5 knockdown caused an over 1-log increase. With the Omicron strain, knockdown of the two E3 ligases resulted in an even greater increase in viral titres (Fig. 5C–E for the Wuhan strain and 5F–H for the Omicron strain). Immunoblotting (IB) confirmed these findings, demonstrating elevated intracellular and secreted N protein levels upon UBR5 or MARCH7 knockdown (Fig. 5I). The knockdown efficiency of MARCHF7 and UBR5 was also confirmed (Fig. 5B). Due to low transfection efficiency in Caco2 cells, we overexpressed UBR5-Myc or MARCH7-Myc in HEK293T-ACE2 cells. Overexpression of UBR5 or MARCH7 significantly decreased viral mRNA levels of the M and E genes in both the Wuhan and Omicron strains, as measured by real-time quantitative polymerase chain reaction (RT-qPCR). IB analysis also showed decreased N protein levels in cells and supernatants, accompanied by a decrease in viral titres of more than 0.5-log. However, co-transfection with increasing amounts of nsp16 significantly abrogated the inhibitory effects of UBR5 and MARCH7 on SARS-CoV-2 replication (Fig. 6A–H and Figure 6—figure supplement 1A–H). To eliminate the potential influence of nsp16’s intrinsic function on experimental outcomes, we overexpressed UBR5 with an inactive HECT domain (UBR5-ΔHECT) or MARCHF7 with a RING domain deletion (MARCHF7-1-542), along with nsp16. These E3 ligase mutants failed to inhibit viral replication. Moreover, a gradual enhancement of nsp16 protein expression did not enhance viral replication. While there was a slight increase in the mRNA levels of M, protein levels of E and N remained unchanged (Figure 6 — figure supplement 2A–H). These results suggest that UBR5 and MARCH7 exhibit antiviral activity against SAR-CoV-2 of UBR5 and MARCHF7 through their enzymatic function, significantly suppressing viral replication by targeting nsp16 for degradation.

Increased levels of nsp16 rescued viral inhibition by UBR5 or MARCHF7’

(A-H) UBR5 or MARCHF7 was transfected in 293T cells stably overexpressed with ACE2, and the increased doses of nsp16-Flag was transfected simultaneously. After 24 h, the cells were infected with Wuhan strains. The mRNA levels of M and E genes of Wuhan strain in the cells (A, D) and E gene in supernatant (B, E) were detected by RT-qPCR, as well as the detection of viral titers in supernatant (C, F). The N protein of the virus and the overexpression efficiency was detected by IB (G, H). Data are representative of three independent experiments and shown as average ±SD (n = 3). Significance was determined by one-way ANOVA, followed by a Tukey multiple comparisons posttest. P > 0.05; **P < 0.01; ***P < 0.001. Figure 6—figure supplement 1 shows data related to infection with Omicron BA.1.

Earlier version: UBR5 and MARCHF7 mediate broad-spectrum degradation of nsp16 variants

Revised version: UBR5 and MARCHF7 mediate nsp16 variant degradation

Earlier version: Given the ongoing emergence of SARS-CoV-2 variants, broad-spectrum antiviral activity remains critical. To investigate whether UBR5 and MARCH7 can mediate broad-spectrum degradation of various nsp16 variants, we aligned the amino acid sequences of nsp16 proteins obtained from the National Center for Biotechnology Information (Figure 6 — figure supplement 3A). Among these, only three variants, XBB.1.9.1, XBB.1.5, and XBB.1.16, exhibited a higher number of mutation sites compared to other variants, which showed only one or two mutations, suggesting a more conserved sequence. Consequently, using mutagenesis techniques, we synthesised the nsp16 sequences for these three variants along with several single-site mutants from other variants. Treatment with the proteasome inhibitor MG132 restored nsp16 protein levels for all variants (Figure 6—figure supplement 3B), confirming that degradation occurs through the UPS. Furthermore, knockdown of either UBR5 or MARCHF7 increased the stability of nsp16 proteins from these variants, albeit to varying degrees, indicating that these nsp16 proteins are sensitive to UBR5- or MARCHF7-mediated degradation (Figure 6 — figure supplement 3C). Taken together, these findings suggest that UBR5 and MARCHF7 may confer broad-spectrum antiviral activity by targeting nsp16 from diverse SARS-CoV-2 for degradation.

Revised version: With the continuous emergence of SARS-CoV-2 variants, broad-spectrum antiviral strategies have become essential. To assess whether UBR5 and MARCHF7 mediate various nsp16 variants, we analysed nsp16 amino acid sequences from the National Center for Biotechnology Information (NCBI) (Figure 6 — figure supplement 3A). Most variants had one or two mutations, except for XBB.1.9.1, XBB.1.5, and XBB.1.16, which contained more mutations, suggesting a conserved sequence. We used mutagenesis to synthesise nsp16 sequences for these three variants, along with several single-site mutants from other variants. Treatment with MG132 restored nsp16 protein levels for all variants (Figure 6 — figure supplement 3B), confirming UPS-mediated degradation.

Added in the revised version: Furthermore, UBR5 or MARCHF7 knockdown stabilised nsp16 proteins from these variants, though to varying degrees (Figure 6 — figure supplement 3C). These results suggest that UBR5 and MARCHF7 contribute to broad-spectrum antiviral activity by targeting nsp16 from diverse SARS-CoV-2 variants for degradation.

SARS-CoV-2 infection reduces UBR5 and MARCHF7 expression

Earlier version: To explore the relationship between SARS-CoV-2 infection and the expression levels of UBR5 and MARCHF7, we analysed changes in their mRNA and protein levels after infection with either the Wuhan or Omicron strains at varying MOIs, using RT-qPCR and IB, respectively. Both mRNA and protein levels of MARCFH7 decreased with increasing viral titres. However, UBR5 expression initially increased at low titres but continuously declined with increasing titres (Figure 6 — figure supplement 4A-4C). Previous studies have reported a role for UBR5 in regulating interferon-γ-mediated pathways, which might explain the observed increase in UBR5 expression at low titres (Wu et al., 2022). To confirm these findings in vivo, we examined UBR5 and MARCHF7 mRNA levels in peripheral blood mononuclear cells from SARS-CoV-2-infected patients with varying disease severity. While UBR5 mRNA levels negatively correlated with disease progression, MARCHF7 levels showed no significant correlation (Figure 6 — figure supplement 4D). These results suggest that SARS-CoV-2 may suppress host antiviral defences mediated by UBR5 and MARCHF7.

Revised version: To investigate the impact of SARS-CoV-2 infection on UBR5 and MARCHF7 expression, we analysed their mRNA and protein levels following infection with the Wuhan and Omicron strains at varying MOIs using RT-qPCR and immunoblotting. MARCHF7 expression decreased consistently with increasing viral titres, whereas UBR5 levels initially rose at low titres before declining as infection progressed (Figure 6—figure supplement 4A-4C). This transient UBR5 upregulation may be linked to its role in interferon-γ-mediated pathways, as previously reported (Wu et al., 2022). To confirm these findings in vivo, we examined UBR5 and MARCHF7 mRNA levels in peripheral blood mononuclear cells from SARS-CoV-2-infected patients with varying disease severity. UBR5 expression negatively correlated with disease progression, whereas MARCHF7 levels showed no significant correlation (Figure 6—figure supplement 4D). These results suggest that SARS-CoV-2 may actively suppress host antiviral defences by downregulating UBR5 and MARCHF7 expression.

UBR5 and MARCHF7 protect mice from SARS-CoV-2 challenge

Earlier version: Given the lack of known activators or inhibitors for E3 ligases, we developed a transient overexpression method for UBR5 and MARCH7 in mice using high-pressure tail vein injection of plasmids (Bonamassa et al., 2011). To assess their in vivo antiviral effects, mice were injected with plasmids encoding UBR5 or MARCH7, followed by SARS-CoV-2 challenge (Fig. 7A). Both MARCHF7 and UBR5 overexpression significantly decreased viral E gene copy numbers in mouse lungs compared to those in the control group, as well as markedly decreased viral titres in lung tissue homogenates (Fig. 7C-D). Additionally, treated mice experienced less weight loss (Fig. 7E). Histopathological analysis of major organs 5 d post-infection revealed that lungs from mice treated with UBR5 or MARCH7 displayed attenuated lesions, including alveolar contraction and pulmonary oedema, compared to those typically observed in control mice (Fig. 7F). These treatments also reduced the abundance of N proteins (Fig. 7G). Collectively, these results suggest that UBR5 or MARCH7 overexpression inhibits SARS-CoV-2 virulence in vivo by promoting nsp16 protein degradation.

Revised version: Since no specific activators or inhibitors for E3 ligases are available, we used a transient overexpression method using high-pressure tail vein injection of plasmids in mice (Bonamassa et al., 2011). To evaluate their antiviral effects in vivo, mice were injected with plasmids encoding UBR5 or MARCHF7, followed by SARS-CoV-2 challenge (Fig. 7A). E3 ligase overexpression significantly reduced viral E gene copy numbers in the lungs and markedly decreased viral titres (Fig. 7C-D). Additionally, treated mice experienced less weight loss compared to controls (Fig. 7E). Histopathological analysis at 5 days post-infection showed that lungs from UBR5-or MARCHF7-treated mice exhibited reduced alveolar contraction and pulmonary oedema compared to the severe lesions observed in control mice (Fig. 7F). Immunoblotting further confirmed lower N protein abundance in treated mice (Fig. 7G). These results suggest that UBR5 and MARCHF7 overexpression suppresses SARS-CoV-2 virulence in vivo by promoting nsp16 degradation.

In a mouse infection model, overexpression of MARCHF7 or UBR5 exerted inhibitory effects on virus.

Earlier version: (A-G) BLAB/C mice were injected with the corresponding plasmids at 40ug/500ul via the high-pressure tail vein, followed by nasal inoculation with 50µl SARS-CoV-2 virus at a dosage of 105.5 TCID50/mL (created using BioRender.com). IB was used to detect the expression of MARCHF7 or UBR5 in the lung tissues (B). Viral RNA loads in mouse lung tissues were detected by measuring the mRNA levels of the E genes by RT-qPCR (C). Lung tissue was collected, homogenized, and the residue was removed by centrifugation to collect the supernatant. The viral titer was then measured using the TCID50 method (D). Mouse body weight was monitored during the experimental period (E). Representative images of H&E staining of lungs of mice with different treatments. Magnification, ×40. Bars, 20 µm (F). The staining of viral N proteins. Magnification, ×63. Bars, 20 µm. n =3 in each group (G). RT-qPCR was used to measure the expression of cytokines and chemokines in the spleens of mice in each group (H). Statistical significance was analyzed using a one-way analysis of variance with Tukey’s multiple comparisons test. (NS, no significance, *p < 0.05, **p < 0.01, ***p < 0.001).

Revised version: (A-G) BLAB/C mice were injected with the corresponding plasmids at 40ug/500ul via the high-pressure tail vein, followed by nasal inoculation with 50µl SARS-CoV-2 virus at a dosage of 105.5 TCID50/mL (created with BioRender.com and the agreement no. is FO27O5X39A). IB was used to detect the expression of MARCHF7 or UBR5 in the lung tissues (B). Viral RNA loads in mouse lung tissues were detected by measuring the mRNA levels of the E genes by RT-qPCR (C). Lung tissue was collected, homogenized, and the residue was removed by centrifugation to collect the supernatant. The viral titer was then measured using the TCID50 method (D). Mouse body weight was monitored during the experimental period (E). Representative images of H&E staining of lungs of mice with different treatments. Magnification, ×40. Bars, 20 µm (F). The staining of viral N proteins. Magnification, ×63. Bars, 20 µm. n =3 in each group (G). RT-qPCR was used to measure the expression of cytokines and chemokines in the spleens of mice in each group (H). Statistical significance was analyzed using a one-way analysis of variance with Tukey’s multiple comparisons test. (NS, no significance, *p < 0.05, **p < 0.01, ***p < 0.001).

Earlier version: Schematic diagram of MARCHF7 and UBR5 ubiquitinate the SARS-CoV-2 non-structural protein nsp16, leading to its degradation via the proteasomal pathway, thereby affecting viral replication (created using BioRender.com).

Revised version: Schematic diagram of MARCHF7 and UBR5 ubiquitinate the SARS-CoV-2 non-structural protein nsp16, leading to its degradation via the proteasomal pathway, thereby affecting viral replication (created with BioRender.com and the agreement no. is FM27O5ZS3U).

Earlier version: SARS-CoV-2 infection can trigger a severe cytokine storm, thought to be responsible for its high mortality rates (Song et al., 2020) (Hojyo et al., 2020). To examine the effect of UBR5 and MARCH7 on inflammatory responses, we measured levels of key cytokines, including interleukin-6, interleukin-1 receptor antagonist, and interleukin-1β, in the spleens of SARS-CoV-2-infected mice (Makaremi et al., 2022) (Hu et al., 2021). Consistent with their antiviral activity, treatment with UBR5 or MARCH7 significantly reduced the production of these cytokines (Fig. 7H).

Revised version: SARS-CoV-2 infection can induce a severe cytokine storm, contributing to high mortality rates (Song et al., 2020) (Hojyo et al., 2020). To assess the impact of UBR5 and MARCHF7 on inflammatory responses, we measured key cytokine levels — including interleukin-6 (IL-6), interleukin-1 receptor antagonist (IL-1RA), and interleukin-1β (IL-1β)—in the spleens of infected mice (Makaremi et al., 2022) (Hu et al., 2021). Consistent with their antiviral activity, UBR5 and MARCHF7 treatment significantly reduced the production of these inflammatory cytokines (Fig. 7H).

Discussion

Earlier version: SARS-CoV-2 nsp16 functions as a 2’O-MTase, catalysing the methylation of the penultimate nucleotide of the viral RNA cap. This modification produces a 5’-RNA cap structure that mimics host RNA, thereby evading host immune detection and responses (Lin et al., 2020) (Balieiro et al., 2022; Russ et al., 2022). High-resolution structures of nsp16 and nsp16-nsp10 heterodimers have been characterised (Rosas-Lemus et al., 2020) (Klima et al., 2022; Lugari et al., 2010), paving the way for antiviral drug development targeting these heterodimers (Balieiro et al., 2022; Melo-Filho et al., 2022). However, the host factors that interact with and regulate nsp16 remain largely unknown.

Revised version: SARS-CoV-2 nsp16 functions as a 2′-O-methyltransferase (2′-O-MTase), catalysing penultimate nucleotide methylation of the viral RNA cap. This modification allows SARS-CoV-2 to mimic host mRNA, evading immune detection and response (Lin et al., 2020) (Balieiro et al., 2022; Russ et al., 2022). High-resolution structures of nsp16 and nsp16-nsp10 heterodimers have been characterised (Rosas-Lemus et al., 2020) (Klima et al., 2022; Lugari et al., 2010), providing a foundation for antiviral drug development targeting these complexes (Balieiro et al., 2022; Melo-Filho et al., 2022). However, host factors that interact with and regulate nsp16 remain largely unknown.

Earlier version: Previously, our group and others have demonstrated the importance of the UPS in modulating SARS-CoV-2 infection (Xu et al., 2022). For example, host E3 ligases, such as RNF5, Cullin4-DDB1-PRPF19, ZNF598, and TRIM7, ubiquitinate and degrade key SARS-CoV-2 proteins, thereby suppressing viral replication (Li et al., 2023; Liang et al., 2022; Maimaitiyiming et al., 2022; Zhang et al., 2024). Conversely, SARS-CoV-2 employs DUBs to counteract these host defences and enhance its replication (Chen et al., 2024; Gao et al., 2024; Gao et al., 2022; Guo et al., 2021). In this study, we identify nsp16 as a short-lived protein with a half-life of only 2–3 h that is targeted for ubiquitination and degradation by the UPS. Using MS and bioinformatics screening, we identified two E3 ligases, UBR5 and MARCHF7, that directly interact with and target nsp16 for ubiquitination and degradation.

Revised version: Previous studies, including work from our group, have highlighted the role of the UPS in modulating SARS-CoV-2 infection (Xu et al., 2022). Several host E3 ligases—such as RNF5, Cullin4-DDB1-PRPF19, ZNF598, and TRIM7 — ubiquitinate and degrade key SARS-CoV-2 proteins, thereby suppressing viral replication (Li et al., 2023; Liang et al., 2022; Maimaitiyiming et al., 2022; Zhang et al., 2024). In contrast, SARS-CoV-2 uses deubiquitinases (DUBs) to counteract these defences and enhance viral replication (Chen et al., 2024; Gao et al., 2024; Gao et al., 2022; Guo et al., 2021). In this study, we identify nsp16 as a short-lived protein with a half-life of only 2 – 3 h, which is targeted for ubiquitination and degradation via the UPS. Using MS and bioinformatics screening, we identified two E3 ligases, UBR5 and MARCHF7, that directly interact with nsp16 and mediate nsp16 ubiquitination and degradation.

Earlier version: UBR5 is a member of the UBR box protein family and contains a HECT domain essential for its E3 ligase activity (Kim et al., 2021). It is known to regulate tumour growth, metastasis (Xiang et al., 2022) (Qiao et al., 2020), and viral infections. For example, URB5 targets Middle East respiratory syndrome coronavirus ORF4b protein for degradation, dampening the host immune response (Zhou et al., 2022). Conversely, UBR5 acts as a cofactor for Zika virus replication by helping TER94/VCP degrade the capsid protein to expose the viral genome to the cytoplasm (Gestuveo et al.). MARCHF7 is a RING E3 Ub ligase that regulates T cell proliferation, neuronal development, inflammasomes, and tumour progression (Zheng, 2021) (Zhang et al., 2016) (Cai et al., 2022) (Zhao et al., 2018). However, its role in viral infections has been largely unexplored, with only one study identifying it as a host protein interacting with the 2C of foot-and-mouth disease virus (Mahajan et al., 2021).

Revised version: UBR5, a member of the UBR box protein family, contains a HECT domain essential for its E3 ligase activity (Kim et al., 2021). UBR5 regulates tumour growth, metastasis (Xiang et al., 2022) (Qiao et al., 2020), and viral infections. For instance, UBR5 targets the ORF4b protein of Middle East respiratory syndrome coronavirus (MERS-CoV) for degradation, reducing its ability to counteract the host immune response (Zhou et al., 2022). Conversely, UBR5 supports Zika virus replication by assisting TER94/VCP in degrading the capsid protein, allowing viral genome release into the cytoplasm (Gestuveo et al.).

Added in the revised version: MARCHF7, a RING E3 ligase, plays a role in T cell proliferation, neuronal development, inflammasome regulation, and tumour progression (Zheng, 2021) (Zhang et al., 2016) (Cai et al., 2022) (Zhao et al., 2018). However, its involvement in viral infections remains largely unexplored, with one study identifying it as a host protein interacting with the 2C protein of foot-and-mouth disease virus (Mahajan et al., 2021).

Earlier version: Our findings reveal that UBR5 and MARCHF7 independently interact with nsp16 and facilitate its degradation, primarily within the cytoplasmic side of the ER. This reflects an evolutionary arms race between viruses, which develop mechanisms to resist host antiviral responses, and hosts, which deploy multiple strategies to counteract the virus. One such strategy is seen with multiple E3 ligases that target the same viral protein for degradation. Functional analyses demonstrate that UBR5 or MARCHF7 knockdown enhanced SARS-CoV-2 replication, whereas their overexpression impaired replication. Furthermore, we found that when the enzymatic activity of these E3 ligases is inactivated, their antiviral effect is abolished, indicating that their antiviral effects are dependent on their E3 ligase activity. Additionally, we observed that in the absence of functional UBR5 or MARCH7, overexpression of nsp16 does not increase viral levels. This suggests that although nsp16 plays an important role, it cannot independently promote viral replication without the involvement of E3 ligases. This underscores the critical function of UBR5 and MARCHF7 as antiviral factors that suppress SARS-CoV-2 replication by promoting nsp16 degradation.

Revised version: Our findings demonstrate that UBR5 and MARCHF7 independently interact with nsp16 and facilitate its degradation, primarily on the cytoplasmic ER. Our findings highlight the ongoing evolutionary battle between viruses, which evolve strategies to evade host defences, and hosts, which deploy multiple mechanisms to counteract viral infections. One such strategy is the use of multiple E3 ligases targeting the same viral protein for degradation. Functional analyses revealed that knocking down UBR5 or MARCHF7 enhanced SARS-CoV-2 replication, whereas their overexpression impaired it. Moreover, when their enzymatic activity was inactivated, the antiviral effect was abolished, confirming that their E3 ligase activity is essential for viral suppression. Notably, nsp16 overexpression alone did not increase viral levels in the absence of functional UBR5 or MARCHF7. This underscores the critical antiviral role of UBR5 and MARCHF7 in restricting SARS-CoV-2 by promoting nsp16 degradation.

Earlier version: Notably, while nsp16 possesses 18 lysine residues, identifying specific ubiquitination sites proved challenging. Based on structural data of nsp16 (Decroly et al., 2011), we constructed truncated nsp16 mutants to assess their degradation profiles. We found that nsp16 could still be degraded to varying degrees (Appendix-figure S2A-B), suggesting that multiple lysine residues serve as recognition sites for E3 ligases. Nonetheless, the results from these truncated mutants regarding the MG132 recovery provided us with valuable insights. For instance, the Δ2-17 mutant, which lacks this segment, showed more than a two-fold increase in nsp16 expression levels. Conversely, the Δ204-212 mutant exhibited significantly reduced recovery after MG132 treatment compared to the wild-type. These observations offer important clues for further investigation into the interaction regions between nsp16 and MARCHF7 or UBR5, as well as potential ubiquitination sites. MS analysis identified K76 of nsp16 as a ubiquitination site (Appendix-figure S2C). However, an nsp16 K76R mutant was still degraded by E3 ligases, although the degree of ubiquitination was reduced compared to that of the wild-type, indicating that K76 might be one of several potential ubiquitination sites (Appendix-figure S2D-E). Therefore, we hypothesise that nsp16 undergoes ubiquitination at multiple sites and may involve non-canonical ubiquitination pathways at non-lysine residues (McClellan et al., 2019). A comprehensive understanding of the ubiquitination landscape of nsp16 requires further investigation. In addition to UBR5 and MARCHF7, other E3 ligases could contribute to nsp16 regulation, which warrants further studies.

Revised version: Due to the presence of 18 lysine residues in nsp16, identifying specific ubiquitination sites was challenging. Using structural data (Decroly et al., 2011), we constructed truncated nsp16 mutants to assess their degradation profiles. Our results showed that nsp16 mutants degradation occurred to varying degrees (Appendix-figure S2A-B), suggesting that multiple lysine residues serve as recognition sites for E3 ligases. Furthermore, MG132 recovery experiments provided valuable insights. The Δ2-17 mutant, which lacks this segment, exhibited more than a two-fold increase in nsp16 expression, whereas the Δ204-212 mutant showed significantly reduced recovery after MG132 treatment compared to the wild-type. These observations offer important clues for further investigation into the interaction regions between nsp16 and MARCHF7 or UBR5, as well as potential ubiquitination sites. MS analysis identified K76 as a ubiquitination site (Appendix-figure S2C). However, the nsp16 K76R mutant still underwent degradation, though its ubiquitination level was reduced compared to that of the wild-type, indicating that K76 is one of several ubiquitination sites (Appendix-figure S2D-E). This suggests that nsp16 may undergo ubiquitination at multiple sites or through non-canonical pathways at non-lysine residues (McClellan et al., 2019). Beyond UBR5 and MARCHF7, other E3 ligases may also contribute to nsp16 regulation, warranting further investigation into additional host factors involved in SARS-CoV-2 restriction.

Earlier version: Currently, no specific activators or inhibitors exist for UBR5 or MARCH7. To validate their antiviral effects in vivo, we employed high-pressure tail vein injection, a method commonly used for gene delivery to the liver (Kamimura et al., 2014). This method has also been shown to facilitate gene delivery to other tissues, including the lungs (Bonamassa et al., 2011). Gene expression was successfully detected in lung tissues. While this method induced mild immune activation, its effect is relatively minimal compared with that of other techniques (Suda et al., 2023). Moreover, studies suggest that mice and non-human primates exhibit higher resistance to innate immune activation than humans, suggesting that our findings were only minimally influenced by immune activation (Raper et al., 2003). However, this highlights that translating new drug targets into therapeutic treatments will require considerable time and effort. For example, human trials involving the injection of naked DNA (e.g., OTC cDNA) have shown variable outcomes. While no severe immune responses were observed in most participants, one case of fatal immune activation was reported (Raper et al., 2003), highlighting the variability of human immune responses and a key limitation in extrapolating findings from animal models to human trials. This variability might also explain why we did not observe a clear trend for MARCHF7 mRNA levels in peripheral blood samples from patients with varying disease severities. The limited number of patients in our study could further contribute to this observation.

Revised version: Currently, no specific activators or inhibitors exist for UBR5 or MARCHF7. To assess their antiviral effects in vivo, we used high-pressure tail vein injection, a gene delivery method primarily targeting the liver (Kamimura et al., 2014) but also effective for lung tissues (Bonamassa et al., 2011). Gene expression was successfully detected in the lungs, and while this method induced mild immune activation, its effect was minimal compared to other techniques (Suda et al., 2023). Notably, mice and non-human primates exhibit greater resistance to innate immune activation than humans (Raper et al., 2003), suggesting that immune activation had little influence on our results. However, translating these findings into therapeutic applications remains challenging. For instance, human trials using naked DNA injections (e.g., OTC cDNA) have shown variable immune responses, with one case of fatal immune activation reported (Raper et al., 2003). This underscores the variability of human immune responses and the challenges in translating findings from animal models to clinical applications. Such variability may also explain the inconsistent trends in MARCHF7 mRNA levels observed in peripheral blood samples from patients with different disease severities. Additionally, the small sample size in our study may have contributed to this observation.

Earlier version: In this study, we identified the host E3 ligases, UBR5 and MARCHF7, as critical regulators of SARS-CoV-2 nsp16 through ubiquitination and subsequent degradation. We demonstrated a negative correlation between the expression levels of these ligases and the degree of infection, underscoring their antiviral roles. Furthermore, both ligases significantly inhibit SARS-CoV-2 replication. Our findings pave the way for the development of novel therapeutic strategies targeting the UPS for the treatment of COVID-19.

Revised version: In conclusion, in this study, we identified the host E3 ligases UBR5 and MARCHF7 as key regulators of SARS-CoV-2 nsp16, facilitating its ubiquitination and degradation. We observed a negative correlation between their expression levels and infection severity, highlighting their antiviral function. Moreover, both ligases effectively suppressed SARS-CoV-2 replication. These findings provide a foundation for developing UPS-targeted therapeutic strategies for COVID-19 treatment.

Materials and Methods

Reagents and antibodies

The drugs used in this study were as follows: MG132 (catalog no. S2619), Bafilomycin A1 (catalog no. S1413), Bortezomib (catalog no. S1013), Carfilzomib (catalog no. S2853) and Vinblastine (catalog no. S4504) were purchased from Selleck (Houston, TX, USA). Cycloheximide (catalog no. 66-81-9) was purchased from Sigma (Saint Louis, MO, USA). The antibodies used for immunoblotting analysis (IB) were as follows: anti-UBR5 mAb (Proteintech, Rosemont, IL, USA, catalog no. 66937-1-Ig), anti-MARCHF7 mAb (Santa Cruz Biotechnology, Dallas, TX, USA, catalog no. sc-166945), SARS-CoV-2 2’-O-ribose Methytransferase Antibody (nsp16) (Cell signaling, Danvers, Massachusetts, USA, catalog no. #70811), anti-Flag mAb (Sigma, catalog no. F1804), anti-tubulin mAb (Abcam, Cambridge, Cambridgeshire, UK, catalog no. ab11323), anti-hemagglutinin (anti-HA) pAb (Invitrogen, catalog no. 71-5500), anti-Myc pAb (Proteintech, Rosemont, IL, USA, catalog no. 16286-1-AP), and anti-GFP mAb (Abcam, catalog no. ab1218). Primary antibody: anti-COX5A pAb (Sangon Biotec, Shanghai, CHN, catalog no. D261450), anti-PDI-mAb (Proteintech, catalog no. 2E6A11) and Human GM130/GOLGA2 Antibody (R&D Systems, Minneapolis, MN, USA, catalog no. AF8199) were used for immunofluorescence, and SARS-CoV-2 nucleocapsid monoclonal antibody (mAb) (GeneTex, Irvine, CA, USA, catalog no. GTX635679) were used for Immunohistochemistry. Fluorescent secondary antibodies: goat anti-Rabbit IgG (H+L) Highly Cross Adsorbed Secondary Antibody, Alexa Fluor Plus 488 (Invitrogen, catalog no. A11001), and goat anti-Rabbit IgG (H+L) Highly Cross Adsorbed Secondary Antibody, Alexa Fluor Plus 568 (catalog no. A-11011) and immunohistochemical secondary antibodies: Streptavidin-Peroxidase Anti-Rabbit IgG kit (catalog no. KIT-9706) were purchased from Invitrogen and maixin (Fuzhou, CHN) respectively.

Cell lines and viruses

The cell lines used in this study include HEK293T cells (catalog no. CRL-11268), Hela cells (catalog no.CRM-CCL-2), Caco2 cells (catalog no. HTB-37). All cell lines were purchased from American Type Culture Collection (ATCC; Manassas, VA, USA) and tested for mycoplasma. In order to conduct virus related experiments, stable cell lines including Caco2-Nint (stably expressing SARS-CoV-2 N gene) and 293T-ACE2 (stably expressing ACE2) were generated in our laboratory. The cells were cultured in a cell incubator containing 5% CO2 at 37°C using Dulbecco’s modified Eagle’s medium (DMEM; HyClone, Logan, UT, USA) containing 10% heat-inactivated fetal calf serum (FCS, GIBCO BRL, Grand Island, NY, USA), 100 mg/ml penicillin, and 100 ug/ml streptomycin to provide nutrition. We expanded SARS-CoV-2 virus-like-particles (Wuhan-Hu-1, MN908947, Clade:19A) with high replication capacity in Caco2-Nint using a biosafety level-2 (BSL-2) cell culture system. The virus infecting the cells include Omicron BA. 1 strain (human/CHN_CVRI-01/2022) and SARS-CoV-2 Wuhan strain (BetaCoV/Beijing/IME-BJ05-2020, GenBank access no. MT291831.1), and mouse adapted SARS-CoV-2/C57MA14 variant (GenBank: OL913104.1) used for experiments in mice were provided by Key Laboratory of Jilin Province for Zoonosis Prevention and Control, and all experiments for virus infection were performed in Biosafety level 3 (BSL-3) cell culture system.

Plasmids

Eukaryotic expression plasmids encoding SARS-CoV-2 nsp protein were provided by Professor Wang Peihui (Shandong University). The nsp16-HA expression vector was constructed by adding HA tag at the C-terminus. UBR5 (Gnen ID: 51366) and its mutants (UBR5-ΔHECT, UBR5-ΔPABC, UBR5-ΔUBR) expression plasmids were constructed using purchased plasmids from Addgene (Watertown, MA, USA) as templates with no tag or a MYC tag at the N-terminus. The cDNA of 293T cells was used as a template to construct MARCHF7(Gene ID: 64844) and its truncated mutant with a Myc tag at the C-terminus. All the above expression plasmids were inserted into the VR1012 vector. For immunofluorescence (IF) and Fluorescence Resonance Energy Transfer (FRET) experiments, pCDNA3.1-YFP vector (catalog no. 13033) and pECFP-C1 vector (catalog no. 6076-1) were purchased from Addgene and BD (Biosciences Clontech), and MARCHF7 or UBR5 was constructed on pECFP-C1, and nsp16 was constructed on pCDNA3.1-YFP. Single point mutants of nsp16 of different virus subtypes were obtained by point mutagenesis. Multi-site combined mutants were synthesized by Sangon Biotec Company. Human ubiquitin protein and its mutants have been previously described. Primers required for PCR were listed in Appendix Table S1.

RNA extraction and real-time quantitative PCR

We used the way of Trizol (Invitrogen, catalog no. 15596018CN) to extract RNA. The RNA was subsequently reverse transcribed using a High-Capacity cDNA Reverse Transcription kit (Applied Biosystems, Carlsbad, CA, USA, catalog no. 4368814). Finally, PrimeScript RT Master Mix (Takara, Shiga, JPN, catalog no. RR036A) and relative real-time primer were used for qPCR on an Mx3005P instrument (Agilent Technologies, Stratagene, La Jolla, CA, USA). Amplification procedure of the target fragment is as follows: predenaturation at 95°C for 2min, denaturation at 95 °C for 30s, annealing at 55 °C for 30s, and extension at 72 °C for 30s, total of 40 cycles. Real-time primer sequences used in this study were shown in appendix table S1.

Immunoblotting analysis (IB)

Cells or supernatant cultured for a certain time were collected, cell precipitates were resuspended with lysis buffer (50 mM Tris–HCl [pH 7.8], 150 mM NaCl, 1.0% NP-40, 5% glycerol and 4 mM EDTA), 4xloading buffer was added (0.08M Tris [pH 6.8], 2.0%SDS, 10% glycerol, 0.1 M dithiothreitol, and 0.2% bromophenol blue), and samples were lysed by heating at 100℃ for 30min. After removal of cell debris by centrifugation, the supernatant was taken to separate proteins of different sizes by SDS-PAGE. The proteins were transferred onto polyvinylidene fluoride (PVDF) membranes, incubated overnight with the indicated primary antibody, and after 1 hour at room temperature with HRP-conjugated secondary antibodies (Jackson Immunoresearch, West Grove, NJ, USA, catalog no. 115-035-062 for anti-mouse and 111-035-045 for anti-rabbit), the proteins were visualized through Ultrasensitive ECL Chemiluminescence Detection Kit (Proteintech, catalog no. B500024).

Stable cell line generation

We purchased the lentiviral vector pLKO.1-puro (catalog no. 8453) from Adggen. ShRNA targeting target genes were designed and synthesized by Sangon Biotech. After annealing, it was inserted into the vector. Lentivirus was generated by transfection with shRNA (cloned in pLKO.1) or control vector, RRE, VSVG and REV, and the cells were screened with 1/1000 puromycin (3 μg/ml, Sigma, catalog no. P8833) after infection 48 hours. The knockdown efficiency was detected by RT-qPCR or IB. The shRNA target sequences used in this study are shown in appendix table S1.

RNAi

Short interfering RNA (siRNA) used in this study were purchased from RiboBio Co. Ltd. (Guangzhou, CHN). The siRNA sequences for knockdown of MARCHF7 or UBR5 are provided in appendix table S1. The siRNA was transfected into the cells by transfection reagent Lipofectamine 2000 (Invitrogen, catalog no. 11668-019), and the corresponding plasmids were transfected 24 hours later using Lipofectamine 3000 Reagent (Invitrogen, Carlsbad, CA, USA, catalog no. L3000-008). The expression of target genes was detected by RT-qPCR or immunoblotting (IB) analysis 48 hours later.

Immunoprecipitation

After 10 hours of MG132 (10 µM) treatment, the cells were harvested, resuspended in 1ml lysis buffer (50 mM Tris-HCl [pH 7.5], 150 mM NaCl, 0.5% NP-40) containing protease inhibitor (Roche, catalog no. 11697498001), placed on the shaker oscillator for 4 hours to lyse, then centrifuged to remove cell debris, and the supernatant was incubated with the antibody and protein G agarose beads (Roche, Basel, Basel-City, CH, catalog no. 11243233001) at 4°C overnight. Protein G agarose beads was washed 6-8 times with wash buffer (20 mM Tris-HCL [pH 7.5], 100 mM NaCl, 0.1 mM EDTA, and 0.05% Tween 20) at 4°C, centrifuged at 800 g for 1 min. The proteins were eluted by adding loading buffer (0.08 M Tris [pH 6.8], 2.0% SDS, 10% glycerol, 0.1 M dithiothreitol, and 0.2% bromophenol blue), the samples were boiled at 100°C for 10min to elate the proteins. The lysates and immunoprecipitations were detected by IB.

Mass spectrometry (MS)

HEK293T were transfected with SARS-CoV-2-nsp16-Flag for 36 hours, and then harvested after 10 hours of treatment with MG132 or DMSO. Proteins were enriched by co-immunoprecipitation assay. And the elutions were analyzed by MS. MS analysis was performed by the national center for protein science (Beijing, CHN).

Immunofluorescence

At 48 hours after transfection, the solution was discarded and cells were washed twice with preheated PBS for 5min each, followed by fixation with 4% paraformaldehyde at 37°C for 10min. After three washes with PBS, permeabilized with 0.2% Triton X-100 for 5 min at 37°C, and then immediately followed by blocking with 10% fetal bovine serum at room temperature for 1 hour. The blocked cells were incubated overnight with the indicated primary antibodies. The next day, after three washes with PBS, the cells were incubated with corresponding secondary antibodies for 1 hour at room temperature in the dark. The nuclei were stained with DAPI (49,6-diamidino-2-phenylindole, Sigma, catalog no. 9542) and then stored in 90% glycerol. The fluorescence was detected by a laser scanning confocal microscope (FV 3000, Olympus, Tokyo, Japan).

Viral infectivity assay

Earlier version: The Caco2 cells were transfected with siRNA by Lipofectamine RNAiMAX Reagent (Invitrogen, catalog no. 13778150) to knockdown UBR5 or MARCHF7 and infected with Wuhan strain (MOI=0.01) or Omicron-BA.1 strain (MOI=0.001) at 24 hours after transfection. After two hours of infection, the Caco2 cells were cultured in fresh medium containing 2% fetal bovine serum for 48 hours. The cells and supernatants were collected.

Revised version: The Caco2 cells were transfected with siRNA by Lipofectamine RNAiMAX Reagent (Invitrogen, catalog no. 13778150) to knockdown UBR5 or MARCHF7 and infected with Wuhan strain (MOI=0.01) or Omicron-BA.1 strain (MOI=0.001) at 24 hours after transfection. After two hours of infection, the Caco2 cells were cultured in fresh medium containing 2% fetal bovine serum for 48 hours. The cells and supernatants were collected. Viral levels were characterized by RT-qPCR and IB analysis of viral structural proteins. The overexpression plasmid UBR5-Myc, MARCHF7-Myc and nsp16-Flag was transfected in 293T-ACE2 by Lipofectamine 3000 Reagent and infected with virus 24 hours later.

Removed in the revised version: Viral levels were characterized by RT-qPCR and IB analysis of viral structural proteins. The overexpression plasmid UBR5-Myc, MARCHF7-Myc and nsp16-Flag was transfected in 293T-ACE2 by Lipofectamine 3000 Reagent and infected with virus 24 hours later.

Mouse lines and infections

BALB/C mice, 6 weeks, were purchased from Charles River Laboratories, Beijing. Mice were cultured in groups according to experimental groups. Our study examined male and female animals, and similar findings are reported for both sexes. To reduce animal suffering, all welfare and experimental procedures were carried out in strict accordance with the Guide for the Care and Use of Laboratory Animals and relevant ethical regulations during the experiment. The mice were randomly divided into 6 groups and injected with 40 µg/500 µl MARCHF7 or UBR5 plasmids via the tail vein at high pressure. Three groups of mice were infected 50ul with SARS-CoV-2 isolates by nasal challenge at a dose of 105.5 TCID50/mL, three groups were not infected with virus as control, and the empty vector injection group was also used as control.

Immunohistochemical analysis

Three mice in each group, a total of 18 mice, were sacrificed after anesthesia. The lungs of mice were collected and fixed in 4% paraformaldehyde fixative. After 2 days, the lungs were progressively dehydrated through different concentrations of ethanol solution, transparently treated, soaked in xylene, and finally embedded in paraffin. Some thin sections obtained through a microtome for hematoxylin and eosin (H&E) staining. The other fraction was incubated with 3% hydrogen peroxide for 5-10 min at room temperature to eliminate endogenous peroxidase activity. SARS-CoV-2 nucleocapsid monoclonal antibody (mAb) (GeneTex, catalog no. GTX635679) and a Streptavidin-Peroxidase Anti-Rabbit IgG kit (Maixin, Fuzhou, CHN, catalog no. KIT-9706) was used to quantify the viral level in lung tissue.

Statistical analysis

The statistical analyses used in the figures have been described in detail in the figure legends. All data were expressed as mean ± standard deviations (SDs). Statistical comparisons were performed using student’s t-test, one-way ANOVA, or repeated measurement ANOVA. The differences were statistically significant as follows: *P < 0.05, **P < 0.01, ***P < 0.001; ns is for no meaning

Ethics statement

Collection of inpatient blood was approved by the Ethics Committee of the First Hospital of Jilin University (21K105-001) in accordance with the guidelines and principles of WMA Declaration of Helsinki and the Department of Health and Human Services Belmont Report. All study participants signed an informed consent form. PBMCS were obtained from 9 COVID-19 patients with mild disease, 6 patients with severe disease, and 5 patients with critical disease (Appendix table S2). All animal experiments were approved by the ethics committee of Research Unit of Key Technologies for Prevention and Control of Virus Zoonoses, Chinese Academy of Medical Sciences, Changchun Veterinary Research Institute, Chineses Academy of Agricultural Sciences (IACUC of AMMS-11-2020-006).